Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).


```python
NAME = ""
```

---

# Sample exercise notebook

Based on exercises at https://exercism.org/tracks/python

## Hello world
### Instructions

The classical introductory exercise. Just say "Hello, World!".

"Hello, World!" is the traditional first program for beginning programming in a new language or environment.

The objectives are simple:

    Modify the provided code so that it produces the string "Hello, World!".
    Run the test suite and make sure that it succeeds.
    Submit your solution and check it at the website.

If everything goes well, you will be ready to fetch your first real exercise.


### Your solution

<!-- grade_id: cell-d824d6a9dd6c5234, points: 0, grade: false, solution: true -->
```python
def hello():
    
    return 'Hello, World!'
```

### Run tests

<!-- grade_id: cell-d3b0ab867157e6c6, points: 0, grade: true, solution: false -->
```python
msg = """
This test expects a return of the string 'Hello, World!'
Did you use print('Hello, World!') by mistake?
"""
assert hello() == "Hello, World!", msg         
```

## Guidos lasagna
### Instructions

You're going to write some code to help you cook a gorgeous lasagna from your favorite cookbook.

You have five tasks, all related to cooking your recipe.

### 1. Define expected bake time in minutes

Define an EXPECTED_BAKE_TIME constant that returns how many minutes the lasagna should bake in the oven. According to your cookbook, the Lasagna should be in the oven for 40 minutes:

#### Your solution

<!-- grade_id: cell-b49c345a32d13e97, points: 0, grade: false, solution: true -->
```python
#TODO: define the 'EXPECTED_BAKE_TIME' constant.

```

#### Run tests

<!-- grade_id: cell-ea6be230b4ea0551, points: 1, grade: true, solution: false -->
```python
failure_msg = 'Expected a constant of EXPECTED_BAKE_TIME with a value of 40.'
assert EXPECTED_BAKE_TIME == 40, failure_msg
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[5], line 2
      1 failure_msg = 'Expected a constant of EXPECTED_BAKE_TIME with a value of 40.'
----> 2 assert EXPECTED_BAKE_TIME == 40, failure_msg

NameError: name 'EXPECTED_BAKE_TIME' is not defined
```

### 2. Calculate remaining bake time in minutes

Implement the bake_time_remaining() function that takes the actual minutes the lasagna has been in the oven as an argument and returns how many minutes the lasagna still needs to bake based on the EXPECTED_BAKE_TIME.

#### Your solution

<!-- grade_id: cell-8b00ef8183a91013, points: 0, grade: false, solution: true -->
```python
#TODO: Remove 'pass' and complete the 'bake_time_remaining()' function below.
def bake_time_remaining():
    """Calculate the bake time remaining.

    :param elapsed_bake_time: int - baking time already elapsed.
    :return: int - remaining bake time (in minutes) derived from 'EXPECTED_BAKE_TIME'.

    Function that takes the actual minutes the lasagna has been in the oven as
    an argument and returns how many minutes the lasagna still needs to bake
    based on the `EXPECTED_BAKE_TIME`.
    """
    
    pass
```

#### Run tests

<!-- grade_id: cell-76836f1fbd7b1d84, points: 1, grade: true, solution: false -->
```python
time = 30
expected = 10
actual_result = bake_time_remaining(time)
failure_msg = f"""
    Called bake_time_remaining({time}).
    The function returned {actual_result}, but the tests
    expected {expected} as the remaining bake time.
"""

assert actual_result == expected, failure_msg
```

### 3. Calculate preparation time in minutes

Implement the preparation_time_in_minutes(number_of_layers) function that takes the number of layers you want to add to the lasagna as an argument and returns how many minutes you would spend making them. Assume each layer takes 2 minutes to prepare.

#### Your solution

<!-- grade_id: cell-2e2653a250dbc0e6, points: 0, grade: false, solution: true -->
```python
#TODO: Define the 'preparation_time_in_minutes()' function below.
# You might also consider using 'PREPARATION_TIME' here, if you have it defined.

```

#### Run tests

<!-- grade_id: cell-da3a9b000f33c1f2, points: 1, grade: true, solution: false -->
```python
layers = 2
actual_result = preparation_time_in_minutes(layers)
expected = 4
failure_msg = f'''
    Called preparation_time_in_minutes({layers}).
    The function returned {actual_result}, but the tests
    expected {expected} as the preparation time.
    '''
assert actual_result == expected, failure_msg
```

### 4. Calculate total elapsed cooking time (prep + bake) in minutes

Implement the elapsed_time_in_minutes(number_of_layers, elapsed_bake_time) function that has two parameters: number_of_layers (the number of layers added to the lasagna) and elapsed_bake_time (the number of minutes the lasagna has been baking in the oven). This function should return the total number of minutes you've been cooking, or the sum of your preparation time and the time the lasagna has already spent baking in the oven.

#### Your solution

<!-- grade_id: cell-5bea1a73ddf30629, points: 0, grade: false, solution: true -->
```python
#TODO: define the 'elapsed_time_in_minutes()' function below.
# Remember to add a docstring (you can copy and then alter the one from bake_time_remaining.)

```

#### Run tests

<!-- grade_id: cell-4918c80791bc128a, points: 1, grade: true, solution: false -->
```python
layers = 3
time = 20
actual_result = elapsed_time_in_minutes(layers, time)
expected = 26
actual_result = elapsed_time_in_minutes(layers, time) 
failure_msg = f'''
    Called elapsed_time_in_minutes({layers}, {time}).
    The function returned {actual_result}, but the tests
    expected {expected} as the elapsed time.
    '''
assert actual_result == expected, failure_msg                               
```

### 5. Update the recipe with notes

Go back through the recipe, adding "notes" in the form of function docstrings.

#### Run tests

<!-- grade_id: cell-9a6a9a3c9b06f3b3, points: 1, grade: true, solution: false -->
```python
assert bake_time_remaining.__doc__ is not None
assert elapsed_time_in_minutes.__doc__ is not None
assert preparation_time_in_minutes.__doc__ is not None
```